---
title: Test
description: Test

---
## Sample exercise

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
key: 00b131ad80
```


`@instructions`

fdgs
<img src="ooups" onerror="alert('XSS')" />
gfdh

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
